<a style="color: black; font-size: x-large ; font-weight: bold">Desarrolle el ejercicio usando las operaciones de Pandas</a> 

```python
import pandas as pd #Carga libreria Pandas
from pandas.testing import assert_frame_equal #LIbrería para evaluar los dataframes
```

## Cree un  diccionario con las notas de 5 alumnos (variable = datos)

> - Clave: alumno
> - Valor: alumno1, alumno2, alumno3, alumno4, alumno5
> - Clave: nota
> - Valor: 3.4, 2.7, 2.2, 4.6, 4.6   

<!-- grade_id: cell-223273bd947f4346, points: 0, grade: false, solution: true -->
```python
datos = pd.DataFrame()
    ### BEGIN SOLUTION
datos = pd.DataFrame(data={"alumno":
          ["alumno1","alumno2","alumno3","alumno4","alumno5"],
          "nota":
          [3.4, 2.7, 2.2, 4.6, 4.6]
          })
### END SOLUTION
```

<!-- grade_id: cell-327564e721f6a023, points: 1.6666666667, grade: true, solution: false -->
```python
pd.testing.assert_frame_equal(datos, datos)
```

## Organizar la tabla con la notas de menor a mayor en una (variable datosDES = )

<!-- grade_id: cell-2ed1fb3fe2d6db53, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
datosDES = datos.sort_values(["nota"], ascending=True)
### END SOLUTION
```

```python
datosDES
```

<!-- grade_id: cell-a5972ffa21ed9e5d, points: 1.6666666667, grade: true, solution: false -->
```python
assert_frame_equal(datosDES, datosDES)
```

##  Agregue una nueva columna _"Grado"_ al final de la tabla, y a todas las filas pongale _"Tercero"_

<!-- grade_id: cell-4c18fd7d4f7dca95, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
datosDES['Grado'] = 'Tercero'
### END SOLUTION
```

```python
datosDES 
```

<!-- grade_id: cell-223095947f9e4dde, points: 1.6666666667, grade: true, solution: false -->
```python
assert_frame_equal(datosDES, datosDES)
```