<a style="color: black; font-size: x-large ; font-weight: bold">Desarrolle el ejercicio usando las operaciones de Pandas</a> 

```python
import pandas as pd #Carga libreria Pandas
from pandas.testing import assert_frame_equal #LIbrería para evaluar los dataframes
```

## Cargue el archivo  

> -url: https://storage.googleapis.com/tf-datasets/titanic/train.csv

<!-- grade_id: cell-223273bd947f4346, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
datos = pd.read_csv("https://storage.googleapis.com/tf-datasets/titanic/train.csv")
### END SOLUTION

```

```python
datos
```

Output:
```
survived     sex   age  n_siblings_spouses  parch     fare   class  \
0           0    male  22.0                   1      0   7.2500   Third   
1           1  female  38.0                   1      0  71.2833   First   
2           1  female  26.0                   0      0   7.9250   Third   
3           1  female  35.0                   1      0  53.1000   First   
4           0    male  28.0                   0      0   8.4583   Third   
..        ...     ...   ...                 ...    ...      ...     ...   
622         0    male  28.0                   0      0  10.5000  Second   
623         0    male  25.0                   0      0   7.0500   Third   
624         1  female  19.0                   0      0  30.0000   First   
625         0  female  28.0                   1      2  23.4500   Third   
626         0    male  32.0                   0      0   7.7500   Third   

        deck  embark_town alone  
0    unknown  Southampton     n  
1          C    Cherbourg     n  
2 
```

<!-- grade_id: cell-327564e721f6a023, points: 1.6666666667, grade: true, solution: false -->
```python
pd.testing.assert_frame_equal(datos, datos)
```

## Una funcion para identificar si una persona dependiendo de su sexo y la clase, puede tener acceso a la popa o al casco del barco

>- Si 'sex:male' y 'class:Third' o 'class:Second' es "Sin acceso a la popa, ni al casco"
>- Si 'sex:female' y 'class:Third' o 'class:Second' es "Con acceso al casco"
>- Si 'sex:male' o 'sex:female' y 'class:First' es "Con acceso a la popa y al casco"

<!-- grade_id: cell-2ed1fb3fe2d6db53, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def acceso(sex,clas):
    if (sex == "male"  and (clas == "Second" or clas == "Third")):
        return 'Sin acceso a la popa, ni al casco'
    elif (sex == "female"  and (clas == "Second" or clas == "Third")):
        return 'Con acceso al casco'
    elif ((sex == "male"  or sex == "female") and clas == "First"):
        return 'Con acceso a la popa y al casco'
### END SOLUTION


```

```python
print(acceso("female","Second")) #Con acceso al casco
```

Output:
```
Con acceso al casco
```

<!-- grade_id: cell-353f52b99f5875e7, points: 1.6666666667, grade: true, solution: false -->
```python
assert acceso("female","Second") == "Con acceso al casco"
### BEGIN HIDDEN TESTS
assert acceso("male","Second") == "Sin acceso a la popa, ni al casco"
assert acceso("male","First") == "Con acceso a la popa y al casco"
assert acceso("female","First") == "Con acceso a la popa y al casco"
### END HIDDEN TESTS
```

##  Aplicar la funcion acceso a una columna  de nombre "acceso" de la tabla  


<!-- grade_id: cell-4c18fd7d4f7dca95, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
datos['Acceso'] = datos.apply(lambda x: acceso(x['sex'],x['class']), axis=1);
### END SOLUTION
```

```python
datos
```

Output:
```
survived     sex   age  n_siblings_spouses  parch     fare   class  \
0           0    male  22.0                   1      0   7.2500   Third   
1           1  female  38.0                   1      0  71.2833   First   
2           1  female  26.0                   0      0   7.9250   Third   
3           1  female  35.0                   1      0  53.1000   First   
4           0    male  28.0                   0      0   8.4583   Third   
..        ...     ...   ...                 ...    ...      ...     ...   
622         0    male  28.0                   0      0  10.5000  Second   
623         0    male  25.0                   0      0   7.0500   Third   
624         1  female  19.0                   0      0  30.0000   First   
625         0  female  28.0                   1      2  23.4500   Third   
626         0    male  32.0                   0      0   7.7500   Third   

        deck  embark_town alone                             Acceso  
0    unknown  Southampton     n  Si
```

<!-- grade_id: cell-223095947f9e4dde, points: 1.6666666667, grade: true, solution: false -->
```python
assert_frame_equal(datos, datos)
```