<a style="color: black; font-size: x-large ; font-weight: bold">Desarrolle el ejercicio usando las operaciones de Pandas</a> 

```python
import pandas as pd #Carga libreria Pandas
from pandas.testing import assert_frame_equal #LIbrería para evaluar los dataframes
```

## Cargue el archivo y guardelo en la variable datos

> -url: https://raw.githubusercontent.com/cgiohidalgo/JupyterHub-nbgrader/main/dataT.csv 

<!-- grade_id: cell-223273bd947f4346, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
datos = pd.read_csv("https://raw.githubusercontent.com/cgiohidalgo/JupyterHub-nbgrader/main/dataT.csv")
### END SOLUTION
```

```python
datos
```

Output:
```
Comida      Color
0  Lechuga      Verde
1   Tomate       Rojo
2     Piña  Amarrillo
3  Naranja    Naranja
```

<!-- grade_id: cell-327564e721f6a023, points: 1.6666666667, grade: true, solution: false -->
```python
pd.testing.assert_frame_equal(datos, datos)
```

## Agregar una columna 

>- nombre columa "Precio"
>- valores [1.964, 2.400, 5.044, 7.48]

<!-- grade_id: cell-2ed1fb3fe2d6db53, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
datos["Precio"] = [1.964, 2.400, 5.044, 7.48]
### END SOLUTION
```

```python
datos
```

Output:
```
Comida      Color  Precio
0  Lechuga      Verde   1.964
1   Tomate       Rojo   2.400
2     Piña  Amarrillo   5.044
3  Naranja    Naranja   7.480
```

<!-- grade_id: cell-a5972ffa21ed9e5d, points: 1.6666666667, grade: true, solution: false -->
```python
assert_frame_equal(datos, datos)
```

##  Modificar el valor de columna con condición

>- Dependiendo de si un producto de la columna 'Precio' es mayor de 2.450 lo clasifique en la 'Categoria' con la palabra 'Costoso'   
>- No puede usar condicionales ni bucles (if-else, while, for) 


<!-- grade_id: cell-4c18fd7d4f7dca95, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
datos.loc[ datos['Precio'] > 2.450 ,"Categoria"] = 'Costoso'
### END SOLUTION
```

```python
datos
```

Output:
```
Comida      Color  Precio Categoria
0  Lechuga      Verde   1.964       NaN
1   Tomate       Rojo   2.400       NaN
2     Piña  Amarrillo   5.044   Costoso
3  Naranja    Naranja   7.480   Costoso
```

<!-- grade_id: cell-223095947f9e4dde, points: 1.6666666667, grade: true, solution: false -->
```python
assert_frame_equal(datos, datos)
```